<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>STACK information item</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<div class="stack-comment">
The purpose of this question is to illustrate an edge case where there are no inputs and no PRTs.
<ul>
<li>Questions with inputs, but no PRTs, can be used for surveys.</li>
<li>Questions without inputs or PRTs can be used as information items, containing computations and (potentially) randomly generated items.</li>
</ul>
</div>
<p>The <em>completed square</em> form of a quadratic is \(a(x-p)^2+q\)</p>
<p>For example \[ {@p2@}={@p1@}\]]]></p></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text/>
    </generalfeedback>
    <defaultgrade>0</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025030700</text>
    </stackversion>
    <questionvariables>
      <text>n1:rand(5)+5;
n2:rand(5)+5;
p1:(x-n1)^2+n2;
p2:expand(p1);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@p2=p1@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <deployedseed>957266243</deployedseed>
    <deployedseed>588971379</deployedseed>
    <deployedseed>2067323478</deployedseed>
  </question>
</quiz>
