<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Enlargement and Similarity: Parsons Proof</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
You have discovered two triangles on the board: the small one, \(ABC\), and the big one, \(AB'C'\), which is an enlargement of the first with centre at \(A\).
</p>
<br>
[[jsxgraph]]
var xmin = {#xmin#};
var xmax = {#xmax#};
var ymin = {#ymin#};
var ymax = {#ymax#};

var board = JXG.JSXGraph.initBoard('jxgbox', {
    boundingbox: [xmin, ymax, xmax, ymin],  
    axis:true,
    showCopyright:false,
    showNavigation:false,
    keepAspectRatio:true
});

// Base triangle coordinates from STACK
var xA = {#xA#};
var yA = {#yA#};
var xB = {#xB#};
var yB = {#yB#};
var xC = {#xC#};
var yC = {#yC#};

// Original triangle ABC
var A = board.create('point', [xA, yA], {name:'A', size:2, fixed:true});
var B = board.create('point', [xB, yB], {name:'B', size:2, fixed:true});
var C = board.create('point', [xC, yC], {name:'C', size:2, fixed:true});
var triABC = board.create('polygon', [A,B,C], {
    borders:{strokeColor:"green", strokeWidth:2},   // green border
    fillColor:"lightgreen",                        // green fill
    fillOpacity:0.4
});

// Scale factor from STACK
var k = {#scale_factor#};

// Image triangle A'B'C' (A' = A)
var xB1 = {#xB1#};
var yB1 = {#yB1#};
var xC1 = {#xC1#};
var yC1 = {#yC1#};

var B1 = board.create('point', [xB1,yB1], {name:"B'", size:2, color:"red", fixed:true});
var C1 = board.create('point', [xC1,yC1], {name:"C'", size:2, color:"red", fixed:true});
var triA1B1C1 = board.create('polygon',[A,B1,C1], {
    borders:{strokeColor:"red", strokeWidth:2, dash:2}, // red border
    fillColor:"pink",                                   // red fill
    fillOpacity:0.3
});

// Optional: show rays from A to dilated points
board.create('line', [A,B1], {straightFirst:true, straightLast:true, strokeColor:"gray", dash:1});
board.create('line', [A,C1], {straightFirst:true, straightLast:true, strokeColor:"gray", dash:1});

[[/jsxgraph]]

<br>
<br>

<p><strong>Theorem:</strong> If \(AB'C'\) is an enlargement of \(ABC\) with centre \(A\) and scale factor \(k\), then \( \triangle ABC \sim \triangle AB'C'\).</p>

<p><em>Drag and drop the steps into the right order:</em></p>
<br>
[[parsons input="ans1" log="true"]]
{# parsons_encode(proof_steps) #}
[[/parsons]]
<p>
[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text/>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/****************************************************************/
stack_include_contrib("prooflib.mac");

/* Random scale factor */
scale_factor : 2 + rand(2);   /* picks either 2 or 3 */

/* Base triangle coordinates (randomised but structured) */
xA : 0;  
yA : 0;

xB : 2 + rand(3);   /* between 2 and 4 */
yB : 0;

xC : 1 + rand(3);   /* between 1 and 3 */
yC : 1 + rand(3);   /* between 1 and 3 */

/* Image triangle after enlargement from A */
xB1 : scale_factor*xB;
yB1 : scale_factor*yB;

xC1 : scale_factor*xC;
yC1 : scale_factor*yC;


/* Find extremes of all points (original + image) */
xmin : min(xA, xB, xC, xB1, xC1) - 1;
xmax : max(xA, xB, xC, xB1, xC1) + 1;
ymin : min(yA, yB, yC, yB1, yC1) - 1;
ymax : max(yA, yB, yC, yB1, yC1) + 1;

/* -------------------------------*/
/****************************************************************/
/* Theorem statement */
thm: "Show that \\(\\triangle ABC \\sim \\triangle AB'C'\\).";

/****************************************************************/
/* Proof steps */ 
proof_steps: [
  ["assumption", "Consider \\(\\triangle ABC\\) and its image \\(\\triangle AB'C'\\) under an enlargement with centre \\(A\\) and scale factor \\(k\\)."],

  ["step1", "Point \\(A\\) is the centre of enlargement, hence \\(A = A'\\)."],

  ["step2", "By definition of enlargement, \\(AB' = k\\cdot AB\\) and \\(AC' = k\\cdot AC\\)."],

  ["step3", "Thus the corresponding side lengths are in proportion: \\(\\tfrac{AB'}{AB} = \\tfrac{AC'}{AC} = k\\)."],

  ["step4", "An enlargement preserves angles at the centre, so \\(\\angle BAC = \\angle B'AC'\\)."],

  ["step5", "Therefore, two pairs of sides are in the same ratio and the included angle is equal (SAS criterion)."],

  ["conclusion", "Hence, \\(\\triangle ABC \\sim \\triangle AB'C'\\)."]
];

/****************************************************************/
/* Correct proof order */
proof_ans: proof("assumption",
    proof_iff("step1","step2"),
    proof_iff("step3","step4"),
    "step5",
    "conclusion"
);


/* Randomize steps for Parsons-style task */
proof_steps: random_permutation(proof_steps);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>1</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>parsons</type>
      <tans>[proof_ans, proof_steps]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>sa: parsons_decode(ans1);
[pd, saa]:proof_assessment(sa, proof_alternatives(proof_ans));</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>pd</sans>
        <tans>0</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>
{@proof_assessment_display(saa, proof_steps)@}
</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>616842617</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>apply(parsons_answer,[proof_ans, proof_steps])</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
