<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Find several x-values on the graph of the trigonometric function y=f(x)=tan(x) or y=cos(x) or y=sin(x)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Consider the graph of the trigonometric function \(y=f(x)={@fx@}\) and the point on the graph at an angle \(x = {@angdeg@}\).</p>
<p>
[[jsxgraph]]
var board = JXG.JSXGraph.initBoard(divid, {
  boundingbox: [-3 * Math.PI - 1, 3, 4 * Math.PI + 1, -3],
  keepaspectratio: false,
  showCopyright: false,
  pan: false
});
var f = board.jc.snippet('{#fx#}', true, 'x', true);
board.create('functiongraph', [f, -3 * Math.PI - 1, 4 * Math.PI + 1]);
var xaxis = board.create('axis', [
  [0, 0],
  [1, 0]
]);
//xaxis.removeAllTicks();
board.create('ticks', [xaxis, Math.PI/4], {
  strokeColor: '#ccc',
  majorHeight: -1,
  minorTicks: 1,
  drawZero: false
});
var yaxis = board.create('axis', [
  [0, 0],
  [0, 1]
]);
var txt = board.create('text', [Math.PI - 0.5, 0.2, "180"], {
  fontSize: 15,
  fixed: true
});
var txt = board.create('text', [2 * Math.PI - 0.5, 0.2, "360"], {
  fontSize: 15,
  fixed: true
});
var txt = board.create('text', [3 * Math.PI - 0.5, 0.2, "540"], {
  fontSize: 15,
  fixed: true
});
var txt = board.create('text', [4 * Math.PI - 0.5, 0.2, "720"], {
  fontSize: 15,
  fixed: true
});
var txt = board.create('text', [-Math.PI -0.9, 0.2, "-180"], {
  fontSize: 15,
  fixed: true
});
var txt = board.create('text', [-2 * Math.PI - 0.9, 0.2, "-360"], {
  fontSize: 15,
  fixed: true
});
var txt = board.create('text', [-3 * Math.PI - 0.9, 0.2, "-540"], {
  fontSize: 15,
  fixed: true
});
var txt = board.create('text', [{#angrad+0.2#}, 0.15, {#angdeg#}], {
  fontSize: 15,
  fixed: true
});
var p1 = board.create('point',[{#angrad#}, 0], {name:'',size:2, fixed: true});
var p2 = board.create('point',[{#angrad#}, {#yval#}], {name:'',size:2, fixed: true});
var vertline = board.create('line', [[{#angrad#}, 0],[{#angrad#}, {#yval#}]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var horline = board.create('line', [[0, {#yval#}],[1, {#yval#}]], {straightFirst: true,straightLast: true,strokeWidth: 2,fixed:true});
[[/jsxgraph]]
<p>Find three other angles in the range \(-540 < x < 720\) which have the same value of \({@fx@}\) as \(x={@angdeg@}\).</p>
</p><p>\(x=\) [[input:ans1]] or [[input:ans2]] or [[input:ans3]]</p>
<p>  [[validation:ans1]] [[validation:ans2]] [[validation:ans3]] [[feedback:prt1]] </p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>
[[ if test="fx = sin(x) "]]
<p>In this case, we have \(x ={@angdeg@} \) and we need to find three other angles that have the same value of \(\sin(x)\) in the given range.</p>

<p>The three other angles can be calculated as follows:</p>
<p>1. \(x_1 = 360 +  {@angdeg@} = {@ta1@}\)</p>
<p>2. \(x_2  = {@angdeg@} - 360 = {@ta2@}\)</p>
<p>3. \(x_3 = 180-{@angdeg@}  = {@ta3@}\)</p>

<p>Therefore, the three angles in the range \(-540 < x < 720\) that have the same value of \(\sin(x)\) as \(x = {@angdeg@}\) are:</p>
<p>\(x = {@ta1@}\) or \(x = {@ta2@}\) or \(x = {@ta3@}\) </p>

[[ elif test="fx = cos(x) "]]
<p>In this case, we have \(x = {@angdeg@}\) and we need to find three other angles that have the same value of \(\cos(x)\) in the given range.</p>

<p>The three other angles can be calculated as follows:</p>
<p>1. \(x_1  = 360 + {@angdeg@} = {@ta1@}\)</p>
<p>2. \(x_2 = {@angdeg@} - 360 = {@ta2@}\)</p>
<p>3. \(x_3  = {@-angdeg@}= {@ta3@}\)</p>

<p>Therefore, the three angles in the range \(-540 < x < 720\) that have the same value of \(\cos(x)\) as \(x = {@angdeg@}\) are:</p>
<p>\(x = {@ta1@}\) or \(x = {@ta2@}\) or \(x = {@ta3@}\)</p>

[[ elif test="fx =tan(x) "]]
<p>In this case, we have \(x = {@angdeg@}\) and we need to find three other angles that have the same value of \(\tan(x)\) in the given range.</p>

<p>The three other angles can be calculated as follows:</p>
<p>1. \(x_1  = 180 + {@angdeg@} = {@ta1@}\)</p>
<p>2. \(x_2 = {@angdeg@} - 180 = {@ta2@}\)</p>
<p>3. \(x_3 = x_1 + 180= {@ta1@} + 180 = {@ta3@}\)</p>

<p>Therefore, the three angles in the range \(-540 < x < 720\) that have the same value of \(\tan(x)\) as \(x = {@angdeg@}\) are:</p>
<p>\(x = {@ta1@}\) or \(x = {@ta2@}\) or \(x = {@ta3@}\)</p>

[[ else ]]
<p>Invalid function. Please check the function \({@fx@}\) and try again.</p>
[[/ if ]]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>angdeg:rand_with_prohib(20,70,[30,45,60]);
angrad:float(angdeg*%pi/180);
angname:string(angdeg);
fx:rand([sin(x),cos(x),tan(x)]);
yval:float(ev(fx,x=angrad));
ta1:if fx=sin(x) then angdeg+360 elseif fx=cos(x) then 360+angdeg else 180+angdeg;
ta2:if fx=sin(x) then angdeg-360 elseif fx=cos(x) then angdeg-360 else -180+angdeg;
ta3:if fx=sin(x) then 180-angdeg elseif fx=cos(x) then -angdeg else angdeg+360;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>the trigonometric function \(y=f(x)={@fx@}\) </p>
<p>Angle point on the graph  \(x = {@angdeg@}\)</p>
<p>The three other angles are:</p>
<p>\(x={@ta1@} \) or \(x={@ta2@}\) or \(x={@ta3@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>setans:{ans1,ans2,ans3};
distans:cardinality(intersect(setans,{angdeg}));
numans:cardinality(setans);
ans12:float(ans1*%pi/180);
ans22:float(ans2*%pi/180);
ans32:float(ans3*%pi/180);
setanstest1:{ev(fx,x=ans12),ev(fx,x=ans22),ev(fx,x=ans32)};
setanstest:member(yval, setanstest1);
setta:{ta1,ta2,ta3};
/*Node 1: Checks if the student setanswers are correct*/
/*Node 2: Checks for the cardinality of the student answers*/
/*Node 3: Checks for the cardinality of the intersection between the student answers and the given angle */
/*Node 4:Checks if the wrong student answer is the sine of the given angle*/
/*Node 5:Checks if the wrong student answer is the cosine of the given angle*/
/*Node 6:Checks if the wrong student answer is the tangent of the given angle*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>setans</sans>
        <tans>setta</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Some of your answers do not have the same value of \({@fx@}\) as \(x={@angdeg@}\).</p>
<p></p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>numans</sans>
        <tans>3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>-</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.5</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>You need three distinct answers. Use the symmetry in the graph to find more possible solutions within the range.</p>
<p></p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>distans</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>-</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You included the angle \(x = {@angdeg@}\), which should not be included as an answer.</p>
<p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>sin(x)</sans>
        <tans>fx</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Hints:</p>
<p>The graphs have a lot of symmetry and you can use it, together with the horizontal line shown in the graph, to find as many solutions as needed.</p>
<p>In the case of \(f(x)=\sin(x)\), the period is 360 degrees so adding 360 to {@angdeg@} as many times as needed should give you a right answer. Also, we can use the symmetry between 0 and 180 degrees. By looking at the graph we see that there are two possible angles (\({@angdeg@}\) and a second one) shown by the intersections between the graph and the horizontal line. Due to the symmetry, we can see that the second angle should be \(x = 180-{@angdeg@} = {@180-angdeg@}\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>cos(x)</sans>
        <tans>fx</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Hints:</p>
<p>The graphs have a lot of symmetry and you can use it, together with the horizontal line shown in the graph, to find as many solutions as needed.</p>
<p>In the case of \(f(x)=\cos(x)\), the period is 360 degrees so adding 360 to {@angdeg@} as many times as needed should give you a right answer. Also, we can use the symmetry between 0 and 360 degrees. By looking at the graph we see that there are two possible angles (\({@angdeg@}\) and a second one) shown by the intersections between the graph and the horizontal line. Due to the symmetry, we can see that the second angle should be \(x = 360-{@angdeg@} = {@360-angdeg@}\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>5</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>tan(x)</sans>
        <tans>fx</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-6-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Hints:</p>
<p>The graphs have a lot of symmetry and you can use it, together with the horizontal line shown in the graph, to find as many solutions as needed.</p>
<p>In the case of \(f(x)=\tan(x)\), the period is 180 degrees so adding 180 to {@angdeg@} as many times as needed should give you the right answers. This can be evidenced in the graph as the intersections between the graph and the horizontal line repeat every 180 degrees.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-6-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1683707094</deployedseed>
    <deployedseed>699269133</deployedseed>
    <deployedseed>1849153945</deployedseed>
    <deployedseed>1263345796</deployedseed>
    <deployedseed>514935983</deployedseed>
    <deployedseed>1861608273</deployedseed>
    <deployedseed>671477541</deployedseed>
    <deployedseed>2031531365</deployedseed>
    <deployedseed>724057093</deployedseed>
    <deployedseed>315951614</deployedseed>
    <deployedseed>958074057</deployedseed>
    <deployedseed>1989793455</deployedseed>
    <deployedseed>1805714053</deployedseed>
    <deployedseed>753474095</deployedseed>
    <deployedseed>1968010891</deployedseed>
    <deployedseed>2088198815</deployedseed>
    <deployedseed>270480830</deployedseed>
    <deployedseed>1678868099</deployedseed>
    <deployedseed>809141923</deployedseed>
    <deployedseed>2050682122</deployedseed>
    <deployedseed>91360578</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
