<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 206867  -->
  <question type="stack">
    <name>
      <text>Plöt</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p></p><p>Give a polynomial (of {@x@} and nothing else) to plot.<br></p><p>[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>/* No wars we are just plotting the student input. */</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text></text>
    </questionnote>
    <questiondescription format="html">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>x^3-x*x+1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ans1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your polynomial as plotted by string replacement, won't survive if you wrap that polynomial in <code>sqrt(...)</code>:</p>

[[jsxgraph]]
[[comment]] Use the [[jsstring]]-block instead of "{#...#}" if there is any risk that 
   the value has something that would need escaping.
   If it is student input then there most definitely is that risk. [[/comment]]
var ex = [[jsstring]]{#ans1#}[[/jsstring]];
ex = ex.replaceAll("^", "**");
var board =  JXG.JSXGraph.initBoard(divid, {boundingbox: [-10, 5, 10, -5], axis: true});

[[comment]]  This is not the way you want to do this, if there is any chance 
   that the student can push arbitrary strings through the validation.  [[/comment]]
var f = (x) => eval(ex);
board.create('functiongraph', [f]);
[[/jsxgraph]]

<p>Your polynomial as plotted by direct JessieCode</p>


[[jsxgraph]]
[[comment]] Use the [[jsstring]]-block instead of "{#...#}" if there is any risk that 
   the value has something that would need escaping.
   If it is student input then there most definitely is that risk. [[/comment]]
var ex = [[jsstring]]{#ans1#}[[/jsstring]];

var board =  JXG.JSXGraph.initBoard(divid, {boundingbox: [-10, 5, 10, -5], axis: true});

[[comment]]  This is safer as we parse the thing to JessieCode which is safe to evaluate [[/comment]]
var f = board.jc.snippet(ex, true, 'x', true);
board.create('functiongraph', [f]);
[[/jsxgraph]]

<p>Finally, a numerical plot for cases where it just is no polynomial.</p>

[[jsxgraph]]
[[comment]] Note that sane authors don't write this type of code inline, but one 
  might use the include logic to inline it anyway.

  Start by generating the points to evaluate. Use the second derivative as the adaptive metric for dropping extraneous points. We do the evaluation using a [[define/]]
 [[/comment]]
[[define
 _minx="-10"
 _maxx="10"
 _minstep="(_maxx-_minx)/500"
 _maxstep="(_maxx-_minx)/10"
 _d2="diff(ans1,x,2)"
 _x="makelist(ev(_x,numer),_x,_minx,_maxx,_minstep)"
 _y="makelist(ev(ans1,x=__x,numer),__x,_x)"
/]]


var board =  JXG.JSXGraph.initBoard(divid, {boundingbox: [-10, 5, 10, -5], axis: true});

board.create('curve', [{#_x#},{#_y#}]);

[[/jsxgraph]]]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>